# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

```python
pd.options.display.max_columns
pd.set_option('display.max_columns', None)
hhcaps
```

Output:
```
State  CMS Certification Number (CCN)*  \
0        AL                            17000   
1        AL                            17008   
2        AL                            17009   
3        AL                            17013   
4        AL                            17014   
...     ...                              ...   
11995    TX                           747978   
11996    TX                           747979   
11997    TX                           747980   
11998    TX                           747981   
11999    TX                           747982   

                                           Provider Name  \
0      BUREAU OF HOME & COMMUNITY SERVICES           ...   
1      JEFFERSON COUNTY HOME CARE                    ...   
2      ALACARE HOME HEALTH & HOSPICE                 ...   
3      GENTIVA HEALTH SERVICES                       ...   
4      AMEDISYS HOME HEALTH                          ...   
...                                                  ...   
11995  EX
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = len(hhcaps.columns)
## Solution goes here
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

Output:
```
0         17000
1         17008
2         17009
3         17013
4         17014
          ...  
11995    747978
11996    747979
11997    747980
11998    747981
11999    747982
Name: CMS Certification Number (CCN)*, Length: 12000, dtype: int64
```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
states = hhcaps.groupby('State')
answer = len(states)
## Solution goes here
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
state_count = states.count()['CMS Certification Number (CCN)*'].sort_values(ascending=False)
answer = str(state_count.index[0])
## Solution goes here
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
states = hhcaps.rename(columns={"Star Rating for health team communicated well with them":"StarCom"}).groupby('State')
starcom_rank = states.agg('mean')['StarCom'].sort_values(ascending=False)
answer = starcom_rank.index[0]
## Solution goes here

```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
hospital_based = hhcaps[hhcaps['Type of Ownership'] == "Hospital Based Program"]

answer = hospital_based.agg('mean')['Star Rating for how patients rated overall care from agency']
## Solution goes here

```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
answer = str(hhcaps.nunique().sort_values().index[0])

## Solution goes here

```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
offerings = hhcaps.select_dtypes(include='bool')
offer_all = offerings[(offerings.all(axis='columns') == True)]

answer = len(offer_all)

## Solution goes here
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
import numpy as np

x = pd.DataFrame(hhcaps['Number of completed Surveys'])
y = x[x['Number of completed Surveys'].str.isnumeric()]
answer = y.agg('median')['Number of completed Surveys']
## Solution goes here
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

```python
def is_stl(city):
    x = ( (city.lower().strip() == "st louis") | (city.lower().strip() == "st. louis") | (city.lower().strip() == "saint louis") )
    return x
```

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
stl = hhcaps[hhcaps.apply(lambda x : is_stl(x['City']), axis = 1)]
print("Number of privders in St. Louis is ",len(stl))
print("Here is the value under the City column for row 6131:  ",hhcaps['City'][6131])
answer = (len(stl) - 1)
raise Exception("There should actually be 24, not 23. You missed one where the city is listed as \"ST. LOUIS\" instead of fully spelling it out.")



## Solution goes here
```

Output:
```
Number of privders in St. Louis is  24
Here is the value under the City column for row 6131:   ST. LOUIS           
Exception                                 Traceback (most recent call last)
<ipython-input-169-4b7d267c3123> in <module>
      3 print("Here is the value under the City column for row 6131:  ",hhcaps['City'][6131])
      4 answer = (len(stl) - 1)
----> 5 raise Exception("There should actually be 24, not 23. You missed one where the city is listed as \"ST. LOUIS\" instead of fully spelling it out.")
      6 
      7 

Exception: There should actually be 24, not 23. You missed one where the city is listed as "ST. LOUIS" instead of fully spelling it out.
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python

```

```python

field1 = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'
field2 = 'Number of completed Surveys'
x = hhcaps[hhcaps[field1].str.isnumeric() & hhcaps[field2].str.isnumeric()].astype({field1: float, field2: int}) 
x['top_count'] = (x[field1]*x[field2]/100)
top_box = x.groupby('State')['top_count',field2].agg('sum')
top_box['percent'] = top_box['top_count'] / top_box[field2] * 100
top_box = top_box.sort_values(by='percent', ascending=False)
answer = top_box.index[0]
print(answer)
## Solution goes here

```

Output:
```
MS
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```