# Week 4 Exercises

Modules 15 through 20 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.


### 15.1 Which of the following is **not** a true statement?

> A) Tuples and Lists can both have items appended to the end of them
>
> B) The elements in both Tuples and Lists have a specific order 
>
> C) Tuples and Lists can both be nested
>
> D) A Tuple can be concatenated to another Tuple using the `+` operator

```python
answer151 = "A"
```

### 15.2 What does this code output?

```python
(a, b) = ('First', 'Second')
(b, a) = (a, b)
print(b)
```

> A) `First`
>
> B) `Second`
>
> C) `('First', 'Second')`
>
> D) _nothing_

```python
answer152 = "A"
```

### 15.3 What is the output of the following code?

```python
(a, b) = (['Paul','Alex','Linda','Kathy','Amit','Samira'], [41, 38, 30, 43, 41, 40, 34, 52])
len(b)
```

> A) _Error_
>
> B) `6`
>
> C) `8`
>
> D) _None_

```python
answer153 = "C"
```

### 16.1 What is the output from the code below?

```python
dxs = [['COPD','CHF','MI'],['DVT','DBS'],['GI','CNC']]
len(dxs[1])
```

> A) _Error_
>
> B) `3`
>
> C) `2`
>
> D) _None_

```python
answer161 = "C"
```

### 16.2 What will the resulting value of `vals` be?

```python
vals = [3.14, 1.41]
vals.append([2.72])
vals = vals + ['i']
```

> A) `[3.14, 1.41, 2.72, ['i']]`
>
> B) `[3.14, 1.41, [2.72], 'i']`
>
> C) `[3.14, 1.41, [2.72], ['i']]`
>
> D) `[3.14, 1.41, [2.72, 'i']`

```python
answer162 = "B"
```

### 16.3 Which of the following statements is **not** correct?


> A) Lists can have other lists inside them
>
> B) When you create a list, you must specify how big it will be
>
> C) A list is an ordered set of objects
>
> D) list.sort() will sort the list in place without creating a copy

```python
answer163 = "B"
```

### 17.1 When you retrieve items from a dictionary in a **for** loop, in what order are they processed?

> A) Ascending order each time
>
> B) An arbitary order that is the same each time as long as the dictionary doesn't change
>
> C) Descending order each time
>
> D) A different order that is completely random each time

```python
answer171 = True
#raise Exception('I recall you telling us that the dictionary would be random as it is not in a particular order. However, per this question, I do not believe D to be the correct answer as the order does not change since as of python 3.7 insertion order is maintained, based on my research into the matter. As such, B would seem to be a more accurate answer, though arbitrary would not seem to be entirely correct either, as it is distinctly the insertion order that is followed. See the code below, which stays in insertion order regardless of alphabetical order') 

subjects = {
    
    "B2": { "PatientID": "E872839", "Age": 32 },
    "H7": { "PatientID": "E989891", "Age": 19 },
    "C3": { "PatientID": "E878198", "Age": 19 },
    "E5": { "PatientID": "E143299", "Age": 3  },
    "F6": { "PatientID": "E123332", "Age": 21 },
    "A1": { "PatientID": "E143291", "Age": 19 },
    "D4": { "PatientID": "E871111", "Age": 21 }
    
}
    
for subject in subjects:
    print(subject)
    
print('\n')

print(subjects)

print('\n')

(subjects)


for subject in subjects:
    print(subject)
```

Output:
```
B2
H7
C3
E5
F6
A1
D4


{'B2': {'PatientID': 'E872839', 'Age': 32}, 'H7': {'PatientID': 'E989891', 'Age': 19}, 'C3': {'PatientID': 'E878198', 'Age': 19}, 'E5': {'PatientID': 'E143299', 'Age': 3}, 'F6': {'PatientID': 'E123332', 'Age': 21}, 'A1': {'PatientID': 'E143291', 'Age': 19}, 'D4': {'PatientID': 'E871111', 'Age': 21}}


B2
H7
C3
E5
F6
A1
D4
```

### 17.2 What data types can be used for dictionary *keys*?

> A) Only strings
>
> B) Strings or numbers
>
> C) Any data type at all
>
> D) Only those data types that are "hashable"

```python
answer172 = "D"
```

### 17.3 Which of the following statements is False?


> A) Dictionaries can contain lists
>
> B) Lists can contain dictionaries
>
> C) All dictionary keys must be strings
>
> D) Dictionaries can contain multiple entries for the same key

```python
answer173 = "C" # though I'm not so sure about D... you can set the same key multiple times, but wouldn't it just overwrite previous assignments, and not store multiple entries?
```

**In the exercises below, don't worry about the `def` and other things around the code you need to write.  If there is a skeleton of code already there, just put your code where it says "your solution here"**

### 20.1 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():

    names = [['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    # To help you, I've provided a blank dictionary named `families`.
    families = {}
    last_names = set()
    # Your code should loop through `names` and do something to build
    # up `families` to be the output dictionary shown in the test
    # cells below.

    ### BEGIN SOLUTION
    for last_name in names:
        last_names.add(last_name[0])
    
    for last_name in last_names:
        families[last_name] = list()
    
    for name in names:
        last_name = name[0]
        families[last_name] = families[last_name] + [name[1]]
        
    ### END SOLUTION

    return families

```

```python
list_to_dict()
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'],
 'Duck': ['Donald', 'Daisy'],
 'Westhus': ['Eric']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict() == answer)
```

### 20.2 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):

    conclusions = []
    
    ### BEGIN SOLUTION
    
    for patient in patients:
        if patient[2] > avg_los[patient[1]]:
            conclusion = "too long"
        elif patient[2] < avg_los[patient[1]]:
            conclusion = "too short"
        elif patient[2] == avg_los[patient[1]]:
            conclusion = "just right"
            
        conclusions = conclusions + [[patient[0] , conclusion]]
    
    ### END SOLUTION
    
    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
[['Boal', 'too short'],
 ['Boal', 'too long'],
 ['Jones', 'just right'],
 ['Ashbury', 'too short']]
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```