# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!

MYTABLE='btabman' + '_data'

```

```python
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd

# Then, create a connection to the same database, using the 
# same user name and password in the `mysql-examples.ipynb` 
# file.  Use conn as the database connection variable.

host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'

conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'paulboal_data')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
import pandas as pd
url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'
dhcs = pd.read_csv(url)
display(dhcs)
### SOLUTION GOES HERE

```

Output:
```
Provider_Name         NPI    CCN  \
0                            Sutter Bay Hospitals  1659439834  50008   
1    PRIME HEALTHCARE SERVICES - GARDEN GROVE LLC  1659538858  50230   
2                          ST MARY MEDICAL CENTER  1669456299  50300   
3                       MADERA COMMUNITY HOSPITAL  1669673646  50568   
4                    Temecula Valley Hospital Inc  1679816201  50775   
..                                            ...         ...    ...   
323                   KAISER FOUNDATION HOSPITALS  1326119967  50071   
324      COUNTY OF LOS ANGELES AUDITOR CONTROLLER  1336154020  50717   
325                    REEDLEY COMMUNITY HOSPITAL  1336167550  50192   
326  PROVIDENCE HEALTH SYSTEM-SOUTHERN CALIFORNIA  1336173269  50235   
327                   KAISER FOUNDATION HOSPITALS  1336294040  50411   

    Medicaid_EP_Hospital_Type           Street_Address           City  \
0        Acute Care Hospitals      CASTRO & DUBOCE STS  SAN FRANCISCO   
1        Acute Care Hospi
```

```python
dhcs.to_sql(MYTABLE, conn, index=False)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,19))
assert(list(dxyz.columns) == ['Provider_Name', 'NPI', 'CCN',
       'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
       'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',
       'Total_payments', 'Last_Program_Year', 'Last_Payment_Year',
       'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount', 'LONGITUDE',
       'LATITUDE', 'FID'])
# Provider_Name	NPI	CCN	Medicaid_EP_Hospital_Type	Street_Address	City	County	State	Zip_Code	Payment_Year_Number	Program_Type	Total_payments	
# Last_Program_Year	Last_Payment_Year	Last_Payment_Criteria	Most_Recent_Disbursement_Amount	LONGITUDE	LATITUDE	FID
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

```python
pop_tbl = pd.read_sql_query("SELECT * FROM population",conn)
dhcs = dhcs.rename(columns={'Zip_Code': 'zipcode'})
pop_tbl['zipcode'] = pd.to_numeric(pop_tbl['zipcode'])
data2 = dhcs.merge(pop_tbl, how='inner')
```

```python
pd.options.display.max_rows
pd.set_option('display.max_rows', 300)
pd.options.display.max_columns
pd.set_option('display.max_columns', 100)
```

```python
#display(data2)
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
import numpy as np

pop_zip = pop_tbl.sort_values(by='population', ascending=False).reset_index()[['zipcode', 'population']]
#display(pop_zip)
zc = pop_zip.zipcode[0]
answer = list(data2[data2['zipcode'] == zc]['NPI'].astype('str'))
```

```python
display(answer)
```

Output:
```
['1194016923']
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

```python
#x = data2[data2['Medicaid_EP_Hospital_Type'] == 'Acute Care Hospitals'][['zipcode','Medicaid_EP_Hospital_Type','Provider_Name','Total_payments', 'population']]
#y = data2[data2['Medicaid_EP_Hospital_Type'] == 'Children\'s Hospitals'][['zipcode','Medicaid_EP_Hospital_Type','Provider_Name','Total_payments', 'population']]
#display(x.merge(y,  right_on='zipcode', left_on='zipcode', how='inner'))
#data2.groupby(["zipcode", "Medicaid_EP_Hospital_Type"]).agg({'Total_payments': 'sum'})
```

```python
Total_payment_sums = dhcs.groupby(['zipcode','Medicaid_EP_Hospital_Type'])['Total_payments'].sum().reset_index()
#display(Total_payment_sums)
```

```python
#part4 = data2[['Provider_Name','zipcode','Medicaid_EP_Hospital_Type','Total_payments']]
#part4['Avg_Pay_per_Capita'] = part4.loc[:,'Total_payments'] / part4.loc[:,'population']
part4 = Total_payment_sums.merge(pop_tbl, how='inner')
#display(part4.round(2))
#display(pop_tbl)
```

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE

#answer = part4.groupby('Medicaid_EP_Hospital_Type')['Avg_Pay_per_Capita'].agg('mean')

part4['Avg_Pay_per_Capita'] = part4.loc[:,'Total_payments'] / part4.loc[:,'population']
answer = part4.groupby('Medicaid_EP_Hospital_Type')['Avg_Pay_per_Capita'].agg('mean').reset_index()

print(answer.round(3))
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals           17974.363
1      Children's Hospitals             111.559
```

```python
print("""	Even though this is just for the purpose of an academic exercise and no one is actually using it, I still feel compelled to add clarifying statements to any data I present (can you tell I've spent a lot of time in biology labs?).
      
      As such, it must be noted that the 'per capita' part of the average pay per capita does not refer to any population of the hospital itself. Instead, the pay per capita is a measurement for each zip code of how many payments were made for a given hospital type within that zip code \(summing all hospitals of a given type within the same zip code rather than calculating them as individual hospitals\) divided by the number of people living in that zip code.
      
      So perhaps it would be best to describe is as the payments-to-hospital-type-in-zip-code-per-person-living-in-zip-code. And then the given average pay per capita is in fact an average of the aformentioned metric across zip codes, with all zip codes being weighted equally, and not adjusting for population.
      
      Trying to decipher the significance of this calculation will likely keep me up at night.
      """
      )
```

Output:
```
Even though this is just for the purpose of an academic exercise and no one is actually using it, I still feel compelled to add clarifying statements to any data I present (can you tell I've spent a lot of time in biology labs?).
      
      As such, it must be noted that the 'per capita' part of the average pay per capita does not refer to any population of the hospital itself. Instead, the pay per capita is a measurement for each zip code of how many payments were made for a given hospital type within that zip code \(summing all hospitals of a given type within the same zip code rather than calculating them as individual hospitals\) divided by the number of people living in that zip code.
      
      So perhaps it would be best to describe is as the payments-to-hospital-type-in-zip-code-per-person-lliving-in-zip-code. And then the given average pay per capita is in fact an average of the aformentioned metric across zip codes, with all zip codes being weighted equally, and not adjus
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

```